## EC Notebook for Lecture 34: Hypothesis Testing

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Wednesday, Nov. 20** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

### The GPA Dataset

In the previous labs and ECs, we have calculated the average GPA for all the courses as below.

```python
import pandas as pd
df = pd.read_csv("gpa.csv")
df["size"] = 0
df["GPA"] = 0

grades = ["A+", "A", "A-", "B+", "B", "B-", "C+", "C", "C-", "D+", "D", "D-", "F"]
points = [4, 4, 3.67, 3.33, 3, 2.67, 2.33, 2, 1.67, 1.33, 1, 0.67, 0]
for i in range(len(grades)):
    df["size"] += df[grades[i]]
    df["GPA"] += df[grades[i]] * points[i]
df["GPA"] /= df["size"]
df.sample(5)
```

## 1. The One-Sample Z Test

Suppose someone claims that the average GPA of all the courses in UIUC is 3.33, and you want to find out whether this number is trustworthy. However, you don't have permission to access the whole dataset except for a random sample of 100 courses. How can you test the given claim?

### z-test for means
Randomly sample 100 courses from the given DataFrame, calculate the p-value of this sample. 

<!-- grade_id: cell-3182d2cfb9079fe2, points: 0, grade: false, solution: true -->
```python
df_sample = ...
p_value = ...


p_value
```

### z-test for proportions

Someone also claims that 26% of courses in the DataFrame have over 50 students. To test this claim, you are going to 
1. create a function to randomly sample 100 classes and return the number of courses whose size is larger than 50.
2. Run the sampling function for 100 times and store the results in a new DataFrame `df_res`.
3. Find the min, max and average of the number.

<!-- grade_id: cell-060c81a4477630d2, points: 0, grade: false, solution: true -->
```python
def sample_class():
    ...

df_res = ...
res_min = ...
res_max = ...
res_avg = ...


print(int(res_min), int(res_max), int(res_avg))
```

<!-- grade_id: cell-621e2757b30ebf38, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Part 1,2,3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(len(df_sample) == 100), "Make sure you sample 100 classes."
assert(p_value > 0.01), "Make sure you calculate the right p value."
assert(sample_class() < 50), "Make sure you create the right function."
assert(len(df_res) == 100), "Make sure you create the right dataframe df_res."
assert(int(res_min) < 20 and int(res_max) > 30 and abs(int(res_avg)-26)<2), "Make sure you create the right dataframe df_class."


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()

```

```python

```

```python

```